Before you turn this problem in, make sure your Kernel works. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

**Step 0:** Run the below cell

```python
import os
```

```python
from file_state_manager import load_state, save_state
```

```python
if os.path.exists("lesson_1_checkpoint_final"):
    load_state(lesson=1, checkpoint="final")
else:
    !mkdir git_folder
    os.chdir("git_folder")
    !git init
    !touch my_abs.py
    code = """def my_abs(x):
        if x < 0:
            return -x
        else:
            return x
    """
    with open("my_abs.py", "w") as f:
        f.write(code)
    !touch README.md
    !git add .
```

**Step 1:** Commit your staged files. Your work from Part 1 has been moved to this notebook for you.

<!-- grade_id: cell-2bc8b425bb7bfaa2, points: 0, grade: false, solution: true -->
```python
# TODO: Commit staged files
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7a23386072c7fb7b, points: 1, grade: true, solution: false -->
```python
s = !git log
assert s[0] != "fatal: your current branch 'master' does not have any commits yet"
```

**Step 2:** Connect to a remote repository

<!-- grade_id: cell-2c9ecf896e08b6f5, points: 0, grade: false, solution: true -->
```python
# TODO: connect to this remote repository: git@github.com:nihaar-c/otter-student-dummy.git
# YOUR CODE HERE
raise NotImplementedError()
remote_v_output = !git remote -v
```

```python
# RUN THIS CELL
repo_url = "git@github.com:nihaar-c/otter-student-dummy.git"
```

Run the below test cell to check your work

<!-- grade_id: cell-b89993a10e58efbb, points: 1, grade: true, solution: false -->
```python
expected_strings = ["origin\t" + repo_url + " (fetch)", "origin\t"+ repo_url + " (push)"]
assert(all(any(expected_string in line for line in remote_v_output) for expected_string in expected_strings))
```

**Step 3:** Push to the remote repository. If it doesn't work, add `--force` to the end of the command

<!-- grade_id: cell-08f0663e364ec220, points: 0, grade: false, solution: true -->
```python
# TODO: Push to remote repository

# YOUR CODE HERE
raise NotImplementedError()
```

Run the below tests to check your work

<!-- grade_id: cell-6955ea02ac8039d9, points: 1, grade: true, solution: false -->
```python
status_output = !git status
expected_strings = ["Your branch is up to date with 'origin/master'.", 'nothing to commit, working tree clean']
assert(all(any(expected_string in line for line in status_output) for expected_string in expected_strings))

log_output = !git log origin/master..master
assert(log_output == [])

!git fetch origin
compare_output = !git rev-parse HEAD
remote_head = !git rev-parse origin/master
assert(compare_output[0] == remote_head[0])
```

```python
files = ["git_folder/my_abs.py", "git_folder/README.md"]
save_state(files, lesson=2, checkpoint="final")
```