Make a new Directory titled "git_folder". Inside this folder, initialize a git repository.



```python
import os

import subprocess
def run_shell_command(command):
    process = subprocess.Popen(command, stdout=subprocess.PIPE, stderr=subprocess.PIPE, shell=True)
    output, error = process.communicate()

    if process.returncode != 0:
        print(f"Error: {error.decode().strip()}")
        return None
    
    return output.decode().strip()
```

<!-- grade_id: cell-d3fce04a3536087a, points: 1, grade: true, solution: false -->
```python
try:
    assert os.path.exists("git_folder")
except Exception:
    raise  Exception("Folder doesn't exist!")
result = None
try:
    os.chdir("git_folder")
    result = !git status
    os.chdir("..")
except Exception:
    raise  Exception("Repository doesn't exist!")
assert result != "fatal: not a git repository (or any of the parent directories): .git"
```

Check the status of your repository. Store your command in the variable `st`.


<!-- grade_id: cell-26ead7539bc8f1bb, points: 0, grade: false, solution: true -->
```python
def check_status():
    os.chdir("git_folder")
    # st = ..
    # YOUR CODE HERE
    raise NotImplementedError()
    os.chdir("..")
    return st
```

<!-- grade_id: cell-81dd646e40847aba, points: 1, grade: true, solution: false -->
```python
assert check_status()[0] == "On branch master"
```

Create a file named `README.md` inside this `git_folder` using git commands

<!-- grade_id: cell-0147f64d5a4b0956, points: 0, grade: false, solution: true -->
```python
def create_readme():
    os.chdir("git_folder")
    # store the command in new_file, ex: new_file = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    os.chdir("..")
    return new_file
create_readme()
```

<!-- grade_id: cell-d3c7abb6bc82c233, points: 1, grade: true, solution: false -->
```python
assert os.path.exists("git_folder/README.md")
```

Now, create another file called `my_abs.py` in the same directory. <br />
Populate it with these contents:

```python
'''
def my_abs(x):
    if x < 0:
        return -x
    else:
        return -x
'''
```

<!-- grade_id: cell-a457fa5080101b07, points: 0, grade: false, solution: true -->
```python
def create_abs():
    os.chdir("git_folder")
    # store the command in new_file, ex: new_file = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    os.chdir("..")
    return new_file
create_abs()
```

<!-- grade_id: cell-b986f6425b60850d, points: 1, grade: true, solution: false -->
```python
assert os.path.exists("git_folder/my_abs.py")
```

<!-- grade_id: cell-7635f2c06627e7d2, points: 1, grade: true, solution: false -->
```python
os.chdir("git_folder")
import my_abs
import importlib
importlib.reload(my_abs)
from my_abs import my_abs
os.chdir("..")
assert my_abs(-4) == my_abs(4)
assert my_abs(0) == 0
assert my_abs(5) == 5
assert my_abs(-10) == my_abs(-10)
```

Stage these two files

```python
os.chdir("git_folder")
status = !git status
os.chdir("..")
flag = 0
for i in status:
    if "my_abs.py" in i or "README.md" in i or "On branch master" in i:
        flag += 1
assert flag == 3
```

```python

```

```python

```