# 3: About Git

<img src="../../img/icon-spill.png" width="200" />

## What is Git?

Git is the most widely used **version control system (VCS)** today.  Any VCS uses a database to record changes to a set of files over time, such as the source files in a software project.  The changes over time are called the history **history**.  Using the database, you can retrieve specific **versions** of files from any point in history.

## Why is Git?

<img src="../../img/it-worked-yesterday.png" width="200" />

One of Git's biggest advantages is to make **collaboration** easier.  By using a branching workflow (which we'll discuss below), multiple developers can work independently on the same codebase.  

Another advantage of any VCS is the ability to **respond to bugs**.  If your app is in production, and a bug is discovered, then you can revert the app to a previous stable while you investigate the cause.  The VCS also helps you discover which **specific change and developer** instroduced the bug.  The intent is not to blame or shame anyone.  Rather, the idea is that the developer who introduced the bug is the one who is most likely to know that part of the codebase and respond to that bug quickly.

## Parts of a Git Project Directory

The contents of a Git repo are stored on your computer as normal files in a normal directory.  In Git terminology, these files are the **working tree**.  You can edit these files freely with your usual text editors and IDEs.  

On your computer, changes to the files are stored in the **local repository** (or local **database**).  Soon, we will learn how to store (or **commit**) new versions of the files; and how to retrieve (or **checkout**) previous versions of the files.  

A small project directory is shown below. 
* The database and metadata are stored in the hidden `.git/` directory.  You can treat this as a black box, and your only interaction with the database and metadata should be via `git` commands.
* `my_program.py` and `README.md` are files you are actively working on.  You definitely want to record these in the local repo.
* `__pycache__` contains intermediate bytecode from running your Python program.  You probably do *not* want to record these in the local repo...

``` Console
~/my_git_project$ ls -la
total 95
drwxr-xr-x  4 gburdell  staff   6 Jan 25 12:33 .
drwx--x--x 43 gburdell  staff  97 Jan 25 12:33 ..
drwxr-xr-x  7 gburdell  staff  10 Jan 25 12:33 .git
-rw-r--r--  1 gburdell  staff  69 Jan 25 12:32 my_program.py
drwxr-xr-x  2 gburdell  staff   3 Jan 25 12:33 __pycache__
-rw-r--r--  1 gburdell  staff  51 Jan 25 12:33 README.md
```

## Tracked vs. Untracked Files

Git does not automatically add files from the working tree to the repo.  Instead, the user must explicitly specify which files are stored in the repo (**tracked**) and which are not (**untracked**).  It's important to consider which files to *not* track.  

In a Python project, you should not track byte code (`*.pyc` and `*.pyo` in `__pycache__` directories) since they will be automatically regenerated whenever someone else runs it with a different operating system or Python version.  Tracking them will just take up unnecessary space in the database. 

For C/C++ and other compiled languages, you should not track build artifacts (`*.o` files, libraries, executables) for a simlar reason.  These artifacts will usually not work on a different operating system.

## Quiz

1. How can you check if a directory in your local machine is a Git repository?
    1. You can see the `.` and `..` directories when you do `ls -la` in your terminal.
    2. There exists a `README.md` file, which is a sign that this is a folder necessarily for collaborative work.
    3. A `.git` repository exists, which stores the database and metadata of the Git repo.
    4. You have to download a terminal plugin to see if your current directory is a Git repo. 

<!-- grade_id: cell-b4952fb1f0ed112d, points: 0, grade: false, solution: true -->
```python
# Input the answer below as a string, such as 'A', 'B', 'C', ...
### BEGIN SOLUTION
a1 = 'C'
### END SOLUTION
```

2. It is considered good practice to track all kinds of files in your Git repo so that your collaborators will have the same access to all the files you have when they want to work on the project further

<!-- grade_id: cell-c45a153b9b97b6ca, points: 0, grade: false, solution: true -->
```python
# Input the answer below as a Python boolean.
### BEGIN SOLUTION
a2 = False
### END SOLUTION
```

<!-- grade_id: cell-97451499bd882d4e, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert a1 == 'C'
assert a2 == False
### END HIDDEN TESTS
```