Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

1. Introduce what a fork is and why it's used
2. Show how to create a fork of the repo on gt-ospo that we can work with
3. Mention we will get to the advanced topics of forking later (how to create a PR and propose the upstream version with the fork changes)

To start at this lesson, we will fork the repo from this link: https://github.com/gt-ospo/oss-training-templates
Fork it directly into whichever directory you're working in right now.

<!-- grade_id: cell-d9a7e2689b6d01f5, points: 5, grade: true, solution: false -->
```python
import os
import shutil

cwd = os.getcwd()
repo_dir = os.path.join(cwd, "oss-training-templates")

# Check if the .git directory exists, which indicates a git repo was created
try:
    assert(os.path.isdir(os.path.join(repo_dir, ".git")))
except AssertionError as e:
    # Clean up the test directory
    if os.path.exists(repo_dir):
        shutil.rmtree(repo_dir)
    print(f"Error: {e}")
```

```python

```