Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Checkout commit that goes back to Lesson 2 start point

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_1"))
```

Output:
```
Restoring Git checkpoint for lesson_2_1...
fatal: not a git repository (or any of the parent directories): .git
File ~/Lesson_2/git_state_manager.py:128, in restore_git_checkpoint(sublesson_name)
    126 for branch in checkpoint["branches"]:
    127     if branch not in existing_branches:
--> 128         subprocess.run(["git", "branch", branch], check=True)
    130 subprocess.run(["git", "checkout", checkpoint["current_branch"]], check=True)
    132 subprocess.run(["git", "reset", "--hard", checkpoint["commit"]], check=True)

File /usr/lib/python3.10/subprocess.py:526, in run(input, capture_output, timeout, check, *popenargs, **kwargs)
    524     retcode = process.poll()
    525     if check and retcode:
--> 526         raise CalledProcessError(retcode, process.args,
    527                                  output=stdout, stderr=stderr)
    528 return CompletedProcess(process.args, retcode, stdout, stderr)

CalledProcessError: Command '['git', 'branch', 'main']' returned non-zero exit
```

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-ccde677cdd141ce4, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_2.json test_1

# YOUR CODE HERE
git checkout d4034cf

```

Output:
```
$ git checkout d4034cf

Note: switching to 'd4034cf'.

You are in 'detached HEAD' state. You can look around, make experimental
changes and commit them, and you can discard any commits you make in this
state without impacting any branches by switching back to a branch.

If you want to create a new branch to retain commits you create, you may
do so (now or later) by using -c with the switch command. Example:

  git switch -c <new-branch-name>

Or undo this operation with:

  git switch -

Turn off this advice by setting config variable advice.detachedHead to false

HEAD is now at d4034cf First commit of my_abs (no try/except yet)

Saved git state for test_1.
```

<!-- grade_id: cell-8a343dbb03e8ae2f, points: 4, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_1")

expected_commit = "d4034cff631ebda366cb68353368ddde7eb25c32"
assert git_state["commit"] == expected_commit, f"Expected commit '{expected_commit}', found '{git_state['commit']}'"
```

<!-- grade_id: cell-d14168a85a6d8772, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_2.json test_2

# YOUR CODE HERE
git checkout main
```

Output:
```
$ git checkout main
Your branch is up to date with 'origin/main'.

Previous HEAD position was d4034cf First commit of my_abs (no try/except yet)
Switched to branch 'main'

Saved git state for test_2.
```

<!-- grade_id: cell-a783d993bd874286, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected commit '{expected_branch}', found '{git_state['branch']}'"
```

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_2"))
```

Output:
```
/home/nkim337/Lesson_2/.git_states/lesson_2_2_checkpoint.json
Saved checkpoint for lesson_2_2.
```

```python

```