Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Create a new branch called "types"

```python
import os
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-9d3de495f1acd9db, points: 0, grade: false, solution: true -->
```python
def create_new_branch():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return command

create_new_branch()
```

<!-- grade_id: cell-ed805ca38d12c788, points: 5, grade: true, solution: false -->
```python
output = !git status
assert("On branch types" in output[0])
```

Change my_abs.py to the contents in the cell below

import math
import numbers

def my_abs(x):
    if isinstance(x, numbers.Real):
        if x < 0:
            return -x
        else:
            return x
    elif isinstance(x, numbers.Complex):
        return math.sqrt(
            x.real ** 2 + x.imag ** 2)
    else:
        return math.nan

Commit the changes to the "types" branch with the message: "Using type checks" and then checkout back to the main branch

<!-- grade_id: cell-889d181cecadc25a, points: 0, grade: false, solution: true -->
```python
def commit_and_checkout():
    # YOUR CODE HERE
    raise NotImplementedError()

    return (commit_command, checkout_command)

commit_and_checkout()
```

<!-- grade_id: cell-3a72b029249a53e7, points: 5, grade: true, solution: false -->
```python
output = !git status
assert("On branch main" in output[0])
```