Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Create a new branch called "types"

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-ecef23ba45da8faf, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_3.json test_1

# YOUR CODE HERE
git checkout -b types

```

Output:
```
$ git checkout -b types

Switched to a new branch 'types'

Saved git state for test_1.
```

<!-- grade_id: cell-f5e3b82ce32eae50, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_3.json", "test_1")

expected_branch = "types"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

Change my_abs.py to the contents in the cell below

```python
import math
import numbers

def my_abs(x):
    if isinstance(x, numbers.Real):
        if x < 0:
            return -x
        else:
            return x
    elif isinstance(x, numbers.Complex):
        return math.sqrt(
            x.real ** 2 + x.imag ** 2)
    else:
        return math.nan
```

Commit the changes to the "types" branch with the message: "Using type checks" and then checkout back to the main branch

<!-- grade_id: cell-da337ea02d35b018, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_3.json test_2

# YOUR CODE HERE
git commit -am "Using type checks"
git checkout main
```

Output:
```
$ git commit -am "Using type checks"
[types dffe0da] Using type checks
 Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>
Your name and email address were configured automatically based
on your username and hostname. Please check that they are accurate.
You can suppress this message by setting them explicitly. Run the
following command and follow the instructions in your editor to edit
your configuration file:

    git config --global --edit

After doing this, you may fix the identity used for this commit with:

    git commit --amend --reset-author

 1 file changed, 6 insertions(+), 2 deletions(-)

$ git checkout main
Your branch is up to date with 'origin/main'.

Switched to branch 'main'

Saved git state for test_2.
```

<!-- grade_id: cell-3a72b029249a53e7, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_3.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

```python

```