Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Create a new branch called "almost-eq"

```python
import os
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-eb8c010f8eab02fa, points: 0, grade: false, solution: true -->
```python
def create_new_branch():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return command

create_new_branch()
```

<!-- grade_id: cell-9144db905bf58c70, points: 3, grade: true, solution: false -->
```python
output = !git status
assert("On branch almost-eq" in output[0])
```

Change my_abs.py to the contents in the cell below

import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan

def my_almost_eq(x, y):
    return my_abs(x - y) < 1e-16

Commit the changes to the "almost-eq" branch with the message: "Added my_almost_eq" and then checkout back to the main branch

<!-- grade_id: cell-af3067d05f32a4c6, points: 0, grade: false, solution: true -->
```python
def commit_and_checkout():
    # YOUR CODE HERE
    raise NotImplementedError()

    return (commit_command, checkout_command)

commit_and_checkout()
```

<!-- grade_id: cell-bf9ee57c9c80ab1f, points: 5, grade: true, solution: false -->
```python
output = !git status
assert("On branch main" in output[0])
```