Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Create a new branch called "almost-eq"

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-c942f54052c18ba6, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_4.json test_1

# YOUR CODE HERE
git checkout -b almost-eq

```

Output:
```
$ git checkout -b almost-eq

Switched to a new branch 'almost-eq'

Saved git state for test_1.
```

<!-- grade_id: cell-9144db905bf58c70, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_4.json", "test_1")

expected_branch = "almost-eq"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

Change my_abs.py to the contents in the cell below

```python
import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan

def my_almost_eq(x, y):
    return my_abs(x - y) < 1e-16
```

Commit the changes to the "almost-eq" branch with the message: "Added my_almost_eq" and then checkout back to the main branch

<!-- grade_id: cell-38d3c97762422e74, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_4.json test_2

# YOUR CODE HERE
git commit -am "Added my_almost_eq"
git checkout main
```

Output:
```
$ git commit -am "Added my_almost_eq"
[almost-eq a799779] Added my_almost_eq
 Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>
Your name and email address were configured automatically based
on your username and hostname. Please check that they are accurate.
You can suppress this message by setting them explicitly. Run the
following command and follow the instructions in your editor to edit
your configuration file:

    git config --global --edit

After doing this, you may fix the identity used for this commit with:

    git commit --amend --reset-author

 1 file changed, 4 insertions(+), 1 deletion(-)

$ git checkout main
Your branch is up to date with 'origin/main'.

Switched to branch 'main'

Saved git state for test_2.
```

<!-- grade_id: cell-bf9ee57c9c80ab1f, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_4.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

```python

```