Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Create a new branch called "almost-eq"

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_3"))
```

<!-- grade_id: cell-c942f54052c18ba6, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_4.json test_1

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9144db905bf58c70, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_4.json", "test_1")

expected_branch = "almost-eq"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

Change my_abs.py to the contents in the cell below

```python
import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan

def my_almost_eq(x, y):
    return my_abs(x - y) < 1e-16
```

Commit the changes to the "almost-eq" branch with the message: "Added my_almost_eq" and then checkout back to the main branch

<!-- grade_id: cell-38d3c97762422e74, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_4.json test_2

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bf9ee57c9c80ab1f, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_4.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_4"))
```

```python

```