Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Perform a standard fast-forward merge from the "main" branch, merging the changes from the "types" branch into "main"

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-e1e3ede63aafd9dc, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_5.json test_1

# YOUR CODE HERE
git merge types

```

Output:
```
$ git merge types
Updating b431e75..dffe0da
Fast-forward
 my_abs.py | 8 ++++++--
 1 file changed, 6 insertions(+), 2 deletions(-)

Saved git state for test_1.
```

<!-- grade_id: cell-10baa3eeaa6910ba, points: 7, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_5.json", "test_1")
last_commit_message = git_state["last_5_commits"].split("\n")[0].split(" ", 1)[1]
expected_message = "Using type checks"
assert last_commit_message == expected_message, f"Expected last commit message '{expected_message}', found '{last_commit_message}'"
```

```python

```