Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Perform a standard fast-forward merge from the "main" branch, merging the changes from the "types" branch into "main"

```python
import os
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-cb31aa6f28ca8161, points: 0, grade: false, solution: true -->
```python
def perform_merge():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return command

perform_merge()
```

<!-- grade_id: cell-0763fc6b1da44bfc, points: 5, grade: true, solution: false -->
```python
output = !git log --oneline --graph --decorate --branches
assert any("HEAD -> main" in line and "types" in line for line in output), f"HEAD -> main is not pointing to types. Output:\n{output}"
```