Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Attempt to perform a merge commit from the "main" branch, merging the changes from the "almost-eq" branch into "main"

NOTE: You can include a "--no-edit" flag in your command to prevent a text editor from popping up

```python
import os
from git_state_manager import save_git_state, load_git_state
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-4f4082907655bc68, points: 0, grade: false, solution: true -->
```python
def perform_merge():
    # YOUR CODE HERE
    command = !git merge --no-edit almost-eq
    
    return command

perform_merge()
```

Output:
```
['Auto-merging my_abs.py',
 'CONFLICT (content): Merge conflict in my_abs.py',
 'Automatic merge failed; fix conflicts and then commit the result.']
```

You may receive a merge conflict message. This means changes from the "types" branch and "almost-eq" branches have conflicting/overlapping changes that you need to manually address.

Go to my_abs.py and fix the changes:

- Everything below "<<<<<<< HEAD" up until "=======" is what currently exists in the branch you're on (main)
- "=======" separates the conflicting changes between HEAD (main) and almost-eq.
- Everything above ">>>>>>> almost-eq" is from the almost-eq branch. This is the code from almost-eq that Git is trying to merge into main.
- Remove the "except" block, then remove these three lines to resolve this merge conflict.

Once the conflicts resolved, add your changes, then commit them with the message: "Resolved merge conflict in my_abs.py"

<!-- grade_id: cell-594dd94b8657d832, points: 0, grade: false, solution: true -->
```python
def add_commit_changes():
    # YOUR CODE HERE
    command = !git commit -am "Resolved merge conflict in my_abs.py"
    
    return command

print(add_commit_changes())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_6.json", "test_1"))
```

Output:
```
['[main f7eb81e] Resolved merge conflict in my_abs.py', ' Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>', 'Your name and email address were configured automatically based', 'on your username and hostname. Please check that they are accurate.', 'You can suppress this message by setting them explicitly. Run the', 'following command and follow the instructions in your editor to edit', 'your configuration file:', '', '    git config --global --edit', '', 'After doing this, you may fix the identity used for this commit with:', '', '    git commit --amend --reset-author', '']
---------------------------------------------------------------------
Saved git state for test_1.
```

<!-- grade_id: cell-7d94860c9d1c46e4, points: 0, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_6.json", "test_1")
last_commit_message = git_state["last_5_commits"].split("\n")[0].split(" ", 1)[1]
expected_message = "Resolved merge conflict in my_abs.py"

assert last_commit_message == expected_message, f"Expected last commit message '{expected_message}', found '{last_commit_message}'"
```

<!-- grade_id: cell-802fc5a9bb55351e, points: 5, grade: true, solution: false -->
```python

```