Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Attempt to perform a merge commit from the "main" branch, merging the changes from the "almost-eq" branch into "main"

NOTE: You can include a "--no-edit" flag in your command to prevent a text editor from popping up

```python
import os
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-4f4082907655bc68, points: 0, grade: false, solution: true -->
```python
def perform_merge():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return command

perform_merge()
```

You may receive a merge conflict message. This means changes from the "types" branch and "almost-eq" branches have conflicting/overlapping changes that you need to manually address.

Go to my_abs.py and fix the changes:

- Everything below "<<<<<<< HEAD" up until "=======" is what currently exists in the branch you're on (main)
- "=======" separates the conflicting changes between HEAD (main) and almost-eq.
- Everything above ">>>>>>> almost-eq" is from the almost-eq branch. This is the code from almost-eq that Git is trying to merge into main.
- Remove the "except" block, then remove these three lines to resolve this merge conflict.

Once the conflicts resolved, add your changes, then commit them with the message: "Resolved merge conflict in my_abs.py"

<!-- grade_id: cell-594dd94b8657d832, points: 0, grade: false, solution: true -->
```python
def add_commit_changes():
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return command

add_commit_changes()
```

<!-- grade_id: cell-7d94860c9d1c46e4, points: 0, grade: true, solution: false -->
```python
output = !git log --oneline --graph --decorate --branches
assert any("HEAD -> main" in line and "Resolved merge conflict in my_abs.py" in line for line in output), f"HEAD -> main is not pointing to types. Output:\n{output}"
```

<!-- grade_id: cell-802fc5a9bb55351e, points: 5, grade: true, solution: false -->
```python

```