Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Push your changes to your repo

```python
import os
from git_state_manager import save_git_state, load_git_state
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-10ec06048db44e1d, points: 0, grade: false, solution: true -->
```python
def perform_push():
    # YOUR CODE HERE
    command = !git push
    
    return command

print(perform_push())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_7.json", "test_1"))
```

Output:
```
['To github.com:natekim27/oss-training-templates.git', '   b431e75..f7eb81e  main -> main']
---------------------------------------------------------------------
Saved git state for test_1.
```

<!-- grade_id: cell-20ca29bf152f6737, points: 0, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_7.json", "test_1")

assert "remote_commit" in git_state, "Remote commit info was not saved. Ensure `save_git_state()` includes remote data."
local_commit = git_state["commit"][:7]
remote_commit = git_state["remote_commit"]

assert local_commit == remote_commit, f"Push failed: local commit '{local_commit}' does not match remote commit '{remote_commit}'"
```

```python

```