Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

## Lesson 2.7: Pushing Changes to the Remote Repo (Fork)

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_6"))
```

### Task 1
Push your changes to the remote fork that you created.

<!-- grade_id: cell-c2af87c2d169c764, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_7.json test_1

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-20ca29bf152f6737, points: 4, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_7.json", "test_1")

assert "remote_commit" in git_state, "Remote commit info was not saved. Ensure `save_git_state()` includes remote data."
local_commit = git_state["commit"][:7]
remote_commit = git_state["remote_commit"]

assert local_commit == remote_commit, f"Push failed: local commit '{local_commit}' does not match remote commit '{remote_commit}'"
```

```python

```