Run the below cell

```python
import os
!mkdir git_folder
os.chdir("git_folder")
!git init
!touch my_abs.py
code = """def my_abs(x):
    if x < 0:
        return -x
    else:
        return x
"""
with open("my_abs.py", "w") as f:
    f.write(code)
!touch README.md
!git add .
```

Now commit your files

<!-- grade_id: cell-7a23386072c7fb7b, points: 1, grade: true, solution: false -->
```python
s = !git log
assert s[0] != "fatal: your current branch 'master' does not have any commits yet"
```

Connect to a remote repository

<!-- grade_id: cell-2c9ecf896e08b6f5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
!git remote add origin git@github.com:nihaar-c/otter-student-dummy.git
### END SOLUTION

remote_v_output = !git remote -v
# TODO: Enter your repo url
repo_url = "git@github.com:nihaar-c/otter-student-dummy.git"
```

<!-- grade_id: cell-b89993a10e58efbb, points: 1, grade: true, solution: false -->
```python
expected_strings = ["origin\t" + repo_url + " (fetch)", "origin\t"+ repo_url + " (push)"]
assert(all(any(expected_string in line for line in remote_v_output) for expected_string in expected_strings))
```

Push to the remote repo

<!-- grade_id: cell-08f0663e364ec220, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
!git push -u origin master
### END SOLUTION
```

<!-- grade_id: cell-6955ea02ac8039d9, points: 1, grade: true, solution: false -->
```python
status_output = !git status
expected_strings = ["Your branch is up to date with 'origin/master'.", 'nothing to commit, working tree clean']
assert(all(any(expected_string in line for line in status_output) for expected_string in expected_strings))

log_output = !git log origin/master..master
assert(log_output == [])

!git fetch origin
compare_output = !git rev-parse HEAD
remote_head = !git rev-parse origin/master
assert(compare_output[0] == remote_head[0])
```