```python
'''
OSS Module 1 - Git Basics

Notes: 
- This notebook will require you to add an SSH key to your Github account to permit commands such as 'git clone'.
- This notebook requires a fresh remote repo to be recreated when doing a full runthrough and when autograding.
- When typing in your commands in the cells with 'TODO's, enter a '!' before each command (ex: !git status), or a '%' for 'cd' commands (ex: %cd ~/git-workflow)
'''
```

```python
import os
import shutil

repo_dir = os.path.expanduser("~/git-workflow")

# Clean up any previous test runs
if os.path.exists(repo_dir):
    shutil.rmtree(repo_dir)
```

```python
'''
Section 4.1: Creating a repo and committing files
'''
```

```python
'''
1. Creating a local repo
'''

# TODO: Enter commands for git init
!mkdir ~/git-workflow
%cd ~/git-workflow
!git init
```

Output:
```
/home/nkim337/git-workflow
[33mhint: Using 'master' as the name for the initial branch. This default branch name[m
[33mhint: is subject to change. To configure the initial branch name to use in all[m
[33mhint: of your new repositories, which will suppress this warning, call:[m
[33mhint: [m
[33mhint: 	git config --global init.defaultBranch <name>[m
[33mhint: [m
[33mhint: Names commonly chosen instead of 'master' are 'main', 'trunk' and[m
[33mhint: 'development'. The just-created branch can be renamed via this command:[m
[33mhint: [m
[33mhint: 	git branch -m <name>[m
Initialized empty Git repository in /home/nkim337/git-workflow/.git/
```

<!-- grade_id: cell-fe2e2c496cd54ed5, points: 5, grade: true, solution: false -->
```python
'''
1. Autograder Test
'''

# Check if the .git directory exists, which indicates a git repo was created
try:
    assert(os.path.isdir(os.path.join(repo_dir, ".git")))
except AssertionError as e:
    # Clean up the test directory
    if os.path.exists(repo_dir):
        shutil.rmtree(repo_dir)
    print(f"Error: {e}")

```

```python
# Create files if not existing in directory yet
files_and_contents = [
    ("README.md", "# Git Workflow\n\nThis is an example repo for the GT OSPO VSIP Summer 2024 Program."),
    ("my_abs.py", "def my_abs(x):\n    if x < 0:\n        return -x\n    else:\n        return x\n")
]

for file_name, content in files_and_contents:
    file_path = os.path.join(repo_dir, file_name)
    if not os.path.exists(file_path):
        with open(file_path, 'w') as file:
            file.write(content)
```

```python
'''
2. Show the state of the working tree
'''

# TODO: Enter command for git status
status_output = !git status
```

<!-- grade_id: cell-9092a4ac46ad9788, points: 5, grade: true, solution: false -->
```python
'''
2. Autograder Test
'''

expected_strings = ['No commits yet', 'Untracked files:', '(use "git add <file>..." to include in what will be committed)', 'README.md', 'my_abs.py']
assert(all(any(expected_string in line for line in status_output) for expected_string in expected_strings))
```

```python
'''
3. Stage new changes
'''

# TODO: Enter command for git add
!git add README.md my_abs.py
```

<!-- grade_id: cell-b790a132d4e99406, points: 5, grade: true, solution: false -->
```python
'''
3. Autograder Test
'''

output = !git status
expected_strings = ['No commits yet', 'Changes to be committed:', '  (use "git rm --cached <file>..." to unstage)', 'README.md', 'my_abs.py']
assert(all(any(expected_string in line for line in output) for expected_string in expected_strings))
```

```python
'''
4. Updating the database
'''

# TODO: Enter command for git commit
!git commit -m "First commit of my_abs (no try/except yet)"
```

Output:
```
[master (root-commit) a3f7f30] First commit of my_abs (no try/except yet)
 Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>
Your name and email address were configured automatically based
on your username and hostname. Please check that they are accurate.
You can suppress this message by setting them explicitly. Run the
following command and follow the instructions in your editor to edit
your configuration file:

    git config --global --edit

After doing this, you may fix the identity used for this commit with:

    git commit --amend --reset-author

 2 files changed, 8 insertions(+)
 create mode 100644 README.md
 create mode 100644 my_abs.py
```

<!-- grade_id: cell-c42c17e33f659aaf, points: 5, grade: true, solution: false -->
```python
'''
4. Autograder Test
'''

output = !git log
expected_strings = ['commit', 'First commit of my_abs (no try/except yet)']
assert(all(any(expected_string in line for line in output) for expected_string in expected_strings))
assert(len(output[0]) == 47)
```

```python
'''
Section 4.2: Working with remotes
'''
```

```python
'''
5. Add and manage remotes
'''

# TODO: Enter commands for git remote
!git remote add origin git@github.com:natekim27/git-workflow.git
remote_v_output = !git remote -v

# TODO: Insert your repo URL here for the autograder test
repo_url = "git@github.com:natekim27/git-workflow.git"
```

<!-- grade_id: cell-4bfad76d2ca25c38, points: 5, grade: true, solution: false -->
```python
'''
5. Autograder Test
'''

expected_strings = ["origin\t" + repo_url + " (fetch)", "origin\t"+ repo_url + " (push)"]
assert(all(any(expected_string in line for line in remote_v_output) for expected_string in expected_strings))
```

```python
'''
6. Upload to the remote repo
'''

# TODO: Enter command for git push
!git push -u origin master
```

Output:
```
Enumerating objects: 4, done.
Counting objects: 100% (4/4), done.
Compressing objects: 100% (4/4), done.
Writing objects: 100% (4/4), 426 bytes | 213.00 KiB/s, done.
Total 4 (delta 0), reused 0 (delta 0), pack-reused 0
To github.com:natekim27/git-workflow.git
 * [new branch]      master -> master
Branch 'master' set up to track remote branch 'master' from 'origin'.
```

<!-- grade_id: cell-9bb07b0e1db8593c, points: 5, grade: true, solution: false -->
```python
'''
6. Autograder Test
'''

status_output = !git status
expected_strings = ["Your branch is up to date with 'origin/master'.", 'nothing to commit, working tree clean']
assert(all(any(expected_string in line for line in status_output) for expected_string in expected_strings))

log_output = !git log origin/master..master
assert(log_output == [])

!git fetch origin
compare_output = !git rev-parse HEAD
remote_head = !git rev-parse origin/master
assert(compare_output[0] == remote_head[0])
```

```python
'''
Section 4.3: Downloading a repo
'''
```

```python
repo_dir = os.path.expanduser("~/somewhere-else")

# Clean up any previous test runs
if os.path.exists(repo_dir):
    shutil.rmtree(repo_dir)
```

```python
'''
7. Download a repo
'''

# TODO: Enter commands for git clone
!mkdir ~/somewhere-else
%cd ~/somewhere-else
!git clone git@github.com:natekim27/git-workflow.git
```

Output:
```
/home/nkim337/somewhere-else
Cloning into 'git-workflow'...
remote: Enumerating objects: 4, done.[K
remote: Counting objects: 100% (4/4), done.[K
remote: Compressing objects: 100% (4/4), done.[K
remote: Total 4 (delta 0), reused 4 (delta 0), pack-reused 0 (from 0)[K
Receiving objects: 100% (4/4), done.
```

<!-- grade_id: cell-26c08752fb7b47d1, points: 5, grade: true, solution: false -->
```python
'''
7. Autograder Test
'''

try:
    assert(os.path.isdir(os.path.join(repo_dir, "git-workflow/.git")))
except AssertionError as e:
    # Clean up the test directory
    if os.path.exists(repo_dir):
        shutil.rmtree(repo_dir)
    print(f"Error: {e}")
```

```python
'''
Section 4.4: Adding more changes
'''
```

```python
%cd ~/git-workflow

expected_content = """import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan
"""
file_path = os.path.expanduser("~/git-workflow/my_abs.py")
current_content = ''
if os.path.exists(file_path):
    with open(file_path, 'r') as file:
        current_content = file.read()

if current_content != expected_content:
    with open(file_path, 'w') as file:
        file.write(expected_content)
```

Output:
```
/home/nkim337/git-workflow
```

```python
'''
8. Show changes in modified files
'''

# TODO: Enter command for git diff
diff_output = !git diff
```

<!-- grade_id: cell-a6db745920e36531, points: 5, grade: true, solution: false -->
```python
'''
8. Autograder Test
'''

expected_strings = ['diff --git a/my_abs.py b/my_abs.py',
                    '--- a/my_abs.py',
                    '+++ b/my_abs.py',
                    '+import math', '+',
                    ' def my_abs(x):',
                    '-    if x < 0:',
                    '-        return -x',
                    '-    else:',
                    '-        return x',
                    '+    try:',
                    '+        if x < 0:',
                    '+            return -x',
                    '+        else:',
                    '+            return x',
                    '+    except TypeError:',
                    '+        return math.nan']

assert(all(any(expected_string in line for line in diff_output) for expected_string in expected_strings))
```

```python
'''
9. Staging and committing & Pushing to remote
'''

# TODO: Enter commands for git commit and git push
!git commit -am "Added try except"
!git push
```

Output:
```
[master 101d6db] Added try except
 Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>
Your name and email address were configured automatically based
on your username and hostname. Please check that they are accurate.
You can suppress this message by setting them explicitly. Run the
following command and follow the instructions in your editor to edit
your configuration file:

    git config --global --edit

After doing this, you may fix the identity used for this commit with:

    git commit --amend --reset-author

 1 file changed, 9 insertions(+), 4 deletions(-)
Enumerating objects: 5, done.
Counting objects: 100% (5/5), done.
Compressing objects: 100% (3/3), done.
Writing objects: 100% (3/3), 398 bytes | 199.00 KiB/s, done.
Total 3 (delta 0), reused 0 (delta 0), pack-reused 0
To github.com:natekim27/git-workflow.git
   a3f7f30..101d6db  master -> master
```

<!-- grade_id: cell-d45b92f7fda30b83, points: 5, grade: true, solution: false -->
```python
'''
9. Autograder Test
'''

status_output = !git status
expected_strings = ["Your branch is up to date with 'origin/master'.", 'nothing to commit, working tree clean']
assert(all(any(expected_string in line for line in status_output) for expected_string in expected_strings))

log_output = !git log origin/master..master
assert(log_output == [])

!git fetch origin
compare_output = !git rev-parse HEAD
remote_head = !git rev-parse origin/master
assert(compare_output[0] == remote_head[0])
```

```python
'''
Section 4.5: Getting new changes from the remote
'''
```

```python
%cd ~/somewhere-else/git-workflow
```

Output:
```
/home/nkim337/somewhere-else/git-workflow
```

```python
'''
10. Download changes from the remote
'''

# TODO: Enter command for git pull
!git pull
```

Output:
```
remote: Enumerating objects: 5, done.[K
remote: Counting objects: 100% (5/5), done.[K
remote: Compressing objects: 100% (3/3), done.[K
remote: Total 3 (delta 0), reused 3 (delta 0), pack-reused 0 (from 0)[K
Unpacking objects: 100% (3/3), 378 bytes | 189.00 KiB/s, done.
From github.com:natekim27/git-workflow
   a3f7f30..101d6db  master     -> origin/master
Updating a3f7f30..101d6db
Fast-forward
 my_abs.py | 13 [32m+++++++++[m[31m----[m
 1 file changed, 9 insertions(+), 4 deletions(-)
```

<!-- grade_id: cell-cb804cc59c6840ea, points: 5, grade: true, solution: false -->
```python
'''
10. Autograder Test
'''

cat_output = !cat my_abs.py

expected_strings = ['import math',
                    'def my_abs(x):',
                    '    try:',
                    '        if x < 0:',
                    '            return -x',
                    '        else:',
                    '            return x',
                    '    except TypeError:',
                    '        return math.nan']

assert(all(any(expected_string in line for line in cat_output) for expected_string in expected_strings))
```

```python

```