TODOs:
1. Introduce what a fork is and why it's used
2. Show how to create a fork of the repo on gt-ospo that we can work with
3. Mention we will get to the advanced topics of forking later (how to create a PR and propose the upstream version with the fork changes)

To start at this lesson, we will fork the repo from this link: https://github.com/gt-ospo/oss-training-templates
Clone your fork directly into whichever directory you're working in right now. Make sure to use SSH to clone, after inserting your Github SSH key into your JupyterHub server.

<!-- grade_id: cell-d9a7e2689b6d01f5, points: 5, grade: true, solution: false -->
```python
import os
import shutil

cwd = os.getcwd()
repo_dir = os.path.join(cwd, "oss-training-templates")

# Check if the .git directory exists, which indicates a git repo was created
try:
    assert(os.path.isdir(os.path.join(repo_dir, ".git")))
except AssertionError as e:
    # Clean up the test directory
    if os.path.exists(repo_dir):
        shutil.rmtree(repo_dir)
    print(f"Error: {e}")
```

```python
from git_state_manager import save_git_checkpoint
os.chdir("oss-training-templates")
print(save_git_checkpoint("lesson_2_1"))
```

```python

```