## Forking a Repository

A **fork** is a copy of another repository that:

- Shares the history of the original repository, where the original repository is called the **upstream**.
- Is linked to the original for the purposes of **pull requests**

A fork allows you to make changes to a project **without affecting the upstream**. When you're ready to propose an update to the upstream project, you can subimt a pull request.

A fork operates exactly like any other repo and can have separate Git tasks done to them. For this lesson, you will create a fork to work with.

----------------------------------------------------------------------------------------------------------
To start at this lesson, we will **fork** the repo from this link: 

https://github.com/gt-ospo/oss-training-templates

**Important:** Clone your fork directly into the folder of the lesson you're working in now. Make sure to use **SSH** to clone, after inserting your Github SSH key into your JupyterHub server.

<!-- grade_id: cell-d9a7e2689b6d01f5, points: 5, grade: true, solution: false -->
```python
import os
import shutil

cwd = os.getcwd()
repo_dir = os.path.join(cwd, "oss-training-templates")

# Check if the .git directory exists, which indicates a git repo was created
try:
    assert(os.path.isdir(os.path.join(repo_dir, ".git")))
except AssertionError as e:
    # Clean up the test directory
    if os.path.exists(repo_dir):
        shutil.rmtree(repo_dir)
    print(f"Error: {e}")

os.chdir("oss-training-templates")
```

**NOTE:** Cells below like this save the state of your git repository as a **checkpoint** to revert back to in case of an error. Cells like this will appear at the end of each sublesson for you to save the git state at the end of it.

Make sure to only save the checkpoint for a sublesson when you've **passed** all the test cells in that respective sublesson.

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_1"))
```

```python

```