Checkout commit that goes back to Lesson 2 start point

```python
import os
from git_state_manager import save_git_state, load_git_state
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-49a97601a036344a, points: 0, grade: false, solution: true -->
```python
def checkout_commit():
    ### BEGIN SOLUTION
    command = !git checkout d4034cf
    ### END SOLUTION
    
    return command

print(checkout_commit())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_2.json", "test_1"))
```

Output:
```
['HEAD is now at d4034cf First commit of my_abs (no try/except yet)']
---------------------------------------------------------------------
Saved git state for test_2.2_1.
```

<!-- grade_id: cell-8a343dbb03e8ae2f, points: 4, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_1")

expected_commit = "d4034cff631ebda366cb68353368ddde7eb25c32"
assert git_state["commit"] == expected_commit, f"Expected commit '{expected_commit}', found '{git_state['commit']}'"
```

<!-- grade_id: cell-7583d824edc2c551, points: 0, grade: false, solution: true -->
```python
def checkout_main():
    ### BEGIN SOLUTION
    command = !git checkout main
    ### END SOLUTION

    return command

print(checkout_main())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_2.json", "test_2"))
```

Output:
```
['Previous HEAD position was d4034cf First commit of my_abs (no try/except yet)', "Switched to branch 'main'", "Your branch is up to date with 'origin/main'."]
---------------------------------------------------------------------
Saved git state for test_2.2_2.
```

<!-- grade_id: cell-a783d993bd874286, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected commit '{expected_branch}', found '{git_state['branch']}'"
```

```python

```