## Lesson 2.2: Checking Out to a Specific Commit

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

**NOTE:** Cells like this below load back the state of your git repository from the **checkpoint**  to revert back to in case of an error. This reverts your git state to the state it was at after finishing the previous sublesson.

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_1"))
```

### Task 1
Checkout to the first commit in the commit history (The commit SHA should start with `d4034cf`)

**NOTES:** 
1. Make sure to remove the `raise notImplementedError()` line for all cells where you input your commands.
2. You can enter in commands in cells in the exact same form that you would in the terminal. No extra syntax needed.

<!-- grade_id: cell-ccde677cdd141ce4, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_2.json test_1

### BEGIN SOLUTION
git checkout d4034cf
### END SOLUTION
```

<!-- grade_id: cell-8a343dbb03e8ae2f, points: 4, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_1")

expected_commit = "d4034cff631ebda366cb68353368ddde7eb25c32"
assert git_state["commit"] == expected_commit, f"Expected commit '{expected_commit}', found '{git_state['commit']}'"
```

### Task 2
Checkout back to the `main` branch

<!-- grade_id: cell-d14168a85a6d8772, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_2.json test_2

### BEGIN SOLUTION
git checkout main
### END SOLUTION
```

<!-- grade_id: cell-a783d993bd874286, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected commit '{expected_branch}', found '{git_state['branch']}'"
```

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_2"))
```

```python

```