Checkout commit that goes back to Lesson 2 start point

```python
import os
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-49a97601a036344a, points: 0, grade: false, solution: true -->
```python
def checkout_commit():
    ### BEGIN SOLUTION
    command = !git checkout d4034cf
    ### END SOLUTION
    
    return command

checkout_commit()
```

Output:
```
["Note: switching to 'd4034cf'.",
 '',
 "You are in 'detached HEAD' state. You can look around, make experimental",
 'changes and commit them, and you can discard any commits you make in this',
 'state without impacting any branches by switching back to a branch.',
 '',
 'If you want to create a new branch to retain commits you create, you may',
 'do so (now or later) by using -c with the switch command. Example:',
 '',
 '  git switch -c <new-branch-name>',
 '',
 'Or undo this operation with:',
 '',
 '  git switch -',
 '',
 'Turn off this advice by setting config variable advice.detachedHead to false',
 '',
 'HEAD is now at d4034cf First commit of my_abs (no try/except yet)']
```

<!-- grade_id: cell-59860761655b88d3, points: 5, grade: true, solution: false -->
```python
output = !git status
assert("HEAD detached at " in output[0])
```

<!-- grade_id: cell-7583d824edc2c551, points: 0, grade: false, solution: true -->
```python
def checkout_main():
    ### BEGIN SOLUTION
    command = !git checkout main
    ### END SOLUTION

    return command

checkout_main()
```

Output:
```
['Previous HEAD position was d4034cf First commit of my_abs (no try/except yet)',
 "Switched to branch 'main'",
 "Your branch is up to date with 'origin/main'."]
```

<!-- grade_id: cell-2f09ad8cf4c3eb3d, points: 5, grade: true, solution: false -->
```python
output = !git status
assert("On branch main" in output[0])
```