Checkout commit that goes back to Lesson 2 start point

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_1"))
```

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-ccde677cdd141ce4, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_2.json test_1

### BEGIN SOLUTION
git checkout d4034cf
### END SOLUTION
```

<!-- grade_id: cell-8a343dbb03e8ae2f, points: 4, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_1")

expected_commit = "d4034cff631ebda366cb68353368ddde7eb25c32"
assert git_state["commit"] == expected_commit, f"Expected commit '{expected_commit}', found '{git_state['commit']}'"
```

<!-- grade_id: cell-d14168a85a6d8772, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_2.json test_2

### BEGIN SOLUTION
git checkout main
### END SOLUTION
```

<!-- grade_id: cell-a783d993bd874286, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_2.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected commit '{expected_branch}', found '{git_state['branch']}'"
```

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_2"))
```

```python

```