Create a new branch called "types"

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_2"))
```

<!-- grade_id: cell-ecef23ba45da8faf, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_3.json test_1

### BEGIN SOLUTION
git checkout -b types
### END SOLUTION
```

<!-- grade_id: cell-f5e3b82ce32eae50, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_3.json", "test_1")

expected_branch = "types"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

Change my_abs.py to the contents in the cell below

```python
import math
import numbers

def my_abs(x):
    if isinstance(x, numbers.Real):
        if x < 0:
            return -x
        else:
            return x
    elif isinstance(x, numbers.Complex):
        return math.sqrt(
            x.real ** 2 + x.imag ** 2)
    else:
        return math.nan
```

Commit the changes to the "types" branch with the message: "Using type checks" and then checkout back to the main branch

<!-- grade_id: cell-da337ea02d35b018, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_3.json test_2

### BEGIN SOLUTION
git commit -am "Using type checks"
git checkout main
### END SOLUTION
```

<!-- grade_id: cell-3a72b029249a53e7, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_3.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_3"))
```

```python

```