Create a new branch called "types"

```python
import os
from git_state_manager import save_git_state, load_git_state
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-9d3de495f1acd9db, points: 0, grade: false, solution: true -->
```python
def create_new_branch():
    ### BEGIN SOLUTION
    command = !git checkout -b types
    ### END SOLUTION
    
    return command

print(create_new_branch())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_3.json", "test_1"))
```

Output:
```
["Switched to a new branch 'types'"]
---------------------------------------------------------------------
Saved git state for test_2.3_1.
```

<!-- grade_id: cell-f5e3b82ce32eae50, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_3.json", "test_1")

expected_branch = "types"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

Change my_abs.py to the contents in the cell below

```python
import math
import numbers

def my_abs(x):
    if isinstance(x, numbers.Real):
        if x < 0:
            return -x
        else:
            return x
    elif isinstance(x, numbers.Complex):
        return math.sqrt(
            x.real ** 2 + x.imag ** 2)
    else:
        return math.nan
```

Commit the changes to the "types" branch with the message: "Using type checks" and then checkout back to the main branch

<!-- grade_id: cell-889d181cecadc25a, points: 0, grade: false, solution: true -->
```python
def commit_and_checkout():
    ### BEGIN SOLUTION
    commit_command = !git commit -am "Using type checks"
    checkout_command = !git checkout main
    ### END SOLUTION

    return (commit_command, checkout_command)

print(commit_and_checkout())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_3.json", "test_2"))
```

Output:
```
(['[types 11cdad8] Using type checks', ' Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>', 'Your name and email address were configured automatically based', 'on your username and hostname. Please check that they are accurate.', 'You can suppress this message by setting them explicitly. Run the', 'following command and follow the instructions in your editor to edit', 'your configuration file:', '', '    git config --global --edit', '', 'After doing this, you may fix the identity used for this commit with:', '', '    git commit --amend --reset-author', '', ' 1 file changed, 6 insertions(+), 2 deletions(-)'], ["Switched to branch 'main'", "Your branch is up to date with 'origin/main'."])
---------------------------------------------------------------------
Saved git state for test_2.3_2.
```

<!-- grade_id: cell-3a72b029249a53e7, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_3.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

```python

```