Create a new branch called "almost-eq"

```python
import os
from git_state_manager import save_git_state, load_git_state
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-eb8c010f8eab02fa, points: 0, grade: false, solution: true -->
```python
def create_new_branch():
    ### BEGIN SOLUTION
    command = !git checkout -b almost-eq
    ### END SOLUTION
    
    return command

print(create_new_branch())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_4.json", "test_1"))
```

Output:
```
["Switched to a new branch 'almost-eq'"]
---------------------------------------------------------------------
Saved git state for test_1.
```

<!-- grade_id: cell-9144db905bf58c70, points: 3, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_4.json", "test_1")

expected_branch = "almost-eq"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

Change my_abs.py to the contents in the cell below

```python
import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan

def my_almost_eq(x, y):
    return my_abs(x - y) < 1e-16
```

Commit the changes to the "almost-eq" branch with the message: "Added my_almost_eq" and then checkout back to the main branch

<!-- grade_id: cell-af3067d05f32a4c6, points: 0, grade: false, solution: true -->
```python
def commit_and_checkout():
    ### BEGIN SOLUTION
    commit_command = !git commit -am "Added my_almost_eq"
    checkout_command = !git checkout main
    ### END SOLUTION

    return (commit_command, checkout_command)

print(commit_and_checkout())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_4.json", "test_2"))
```

Output:
```
(['[almost-eq 528c804] Added my_almost_eq', ' Committer: Nathaniel E Kim <nkim337@horizons-research.cc.gatech.edu>', 'Your name and email address were configured automatically based', 'on your username and hostname. Please check that they are accurate.', 'You can suppress this message by setting them explicitly. Run the', 'following command and follow the instructions in your editor to edit', 'your configuration file:', '', '    git config --global --edit', '', 'After doing this, you may fix the identity used for this commit with:', '', '    git commit --amend --reset-author', '', ' 1 file changed, 4 insertions(+), 1 deletion(-)'], ["Switched to branch 'main'", "Your branch is up to date with 'origin/main'."])
---------------------------------------------------------------------
Saved git state for test_2.
```

<!-- grade_id: cell-bf9ee57c9c80ab1f, points: 5, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_4.json", "test_2")

expected_branch = "main"
assert git_state["branch"] == expected_branch, f"Expected branch: '{expected_branch}', found: '{git_state['branch']}'"
```

```python

```