Perform a standard fast-forward merge from the "main" branch, merging the changes from the "types" branch into "main"

```python
import os
from git_state_manager import load_git_state
%load_ext git_magics
os.chdir("oss-training-templates")
```

```python
from git_state_manager import restore_git_checkpoint
print(restore_git_checkpoint("lesson_2_4"))
```

<!-- grade_id: cell-e1e3ede63aafd9dc, points: 0, grade: false, solution: true -->
```python
%%save_git_state git_state_2_5.json test_1

### BEGIN SOLUTION
git merge types
### END SOLUTION
```

<!-- grade_id: cell-10baa3eeaa6910ba, points: 7, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_5.json", "test_1")
last_commit_message = git_state["last_5_commits"].split("\n")[0].split(" ", 1)[1]
expected_message = "Using type checks"
assert last_commit_message == expected_message, f"Expected last commit message '{expected_message}', found '{last_commit_message}'"
```

```python
from git_state_manager import save_git_checkpoint
print(save_git_checkpoint("lesson_2_5"))
```

```python

```