Perform a standard fast-forward merge from the "main" branch, merging the changes from the "types" branch into "main"

```python
import os
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-cb31aa6f28ca8161, points: 0, grade: false, solution: true -->
```python
def perform_merge():
    ### BEGIN SOLUTION
    command = !git merge types
    ### END SOLUTION
    
    return command

perform_merge()
```

Output:
```
['Updating b431e75..ed07d88',
 'Fast-forward',
 ' my_abs.py | 8 ++++++--',
 ' 1 file changed, 6 insertions(+), 2 deletions(-)']
```

<!-- grade_id: cell-0763fc6b1da44bfc, points: 5, grade: true, solution: false -->
```python
output = !git log --oneline --graph --decorate --branches
assert any("HEAD -> main" in line and "types" in line for line in output), f"HEAD -> main is not pointing to types. Output:\n{output}"
```