Attempt to perform a merge commit from the "main" branch, merging the changes from the "almost-eq" branch into "main"

NOTE: You can include a "--no-edit" flag in your command to prevent a text editor from popping up

```python
import os
from git_state_manager import save_git_state, load_git_state
os.chdir("oss-training-templates")
```

<!-- grade_id: cell-4f4082907655bc68, points: 0, grade: false, solution: true -->
```python
def perform_merge():
    ### BEGIN SOLUTION
    command = !git merge --no-edit almost-eq
    ### END SOLUTION
    
    return command

perform_merge()
```

Output:
```
['Auto-merging my_abs.py', 'CONFLICT (content): Merge conflict in my_abs.py', 'Automatic merge failed; fix conflicts and then commit the result.']
---------------------------------------------------------------------
Saved git state for test_1.
```

You may receive a merge conflict message. This means changes from the "types" branch and "almost-eq" branches have conflicting/overlapping changes that you need to manually address.

Go to my_abs.py and fix the changes:

- Everything below "<<<<<<< HEAD" up until "=======" is what currently exists in the branch you're on (main)
- "=======" separates the conflicting changes between HEAD (main) and almost-eq.
- Everything above ">>>>>>> almost-eq" is from the almost-eq branch. This is the code from almost-eq that Git is trying to merge into main.
- Remove the "except" block, then remove these three lines to resolve this merge conflict.

Once the conflicts resolved, add your changes, then commit them with the message: "Resolved merge conflict in my_abs.py"

<!-- grade_id: cell-594dd94b8657d832, points: 0, grade: false, solution: true -->
```python
def add_commit_changes():
    ### BEGIN SOLUTION
    command = !git commit -am "Resolved merge conflict in my_abs.py"
    ### END SOLUTION
    
    return command

print(add_commit_changes())
print("---------------------------------------------------------------------\n" + save_git_state("git_state_2_6.json", "test_1"))
```

Output:
```
['On branch main', "Your branch is ahead of 'origin/main' by 3 commits.", '  (use "git push" to publish your local commits)', '', 'Untracked files:', '  (use "git add <file>..." to include in what will be committed)', '\t.ipynb_checkpoints/', '', 'nothing added to commit but untracked files present (use "git add" to track)']
---------------------------------------------------------------------
Saved git state for test_1.
```

<!-- grade_id: cell-7d94860c9d1c46e4, points: 0, grade: true, solution: false -->
```python
git_state = load_git_state("git_state_2_6.json", "test_1")
last_commit_message = git_state["last_5_commits"].split("\n")[0].split(" ", 1)[1]
expected_message = "Resolved merge conflict in my_abs.py"

assert last_commit_message == expected_message, f"Expected last commit message '{expected_message}', found '{last_commit_message}'"
```

<!-- grade_id: cell-802fc5a9bb55351e, points: 5, grade: true, solution: false -->
```python

```